---
title: fdsfsfs
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise
key: 0beb5b0a4c
xp: 100
```

fdsffssfsfsfs df sf sfs dfs fsdf sf sdf s

`@instructions`
fdsffssfsfsfs df sf sfs dfs fsdf sf sdf s

`@hint`
fdsffssfsfsfs df sf sfs dfs fsdf sf sdf s

`@pre_exercise_code`
```{r}
my_fun <- function() { print('----') }
```

`@sample_code`
```{r}
dasdad
```

`@solution`
```{r}
my_fun <- function() { print('----') }
```

`@sct`
```{r}
ex() %>% check_error()
```
