---
title: 'Motivation for regression'
description: 'Chapter description goes here.'
free_preview: true
---

## gfdgd

```yaml
type: NormalExercise
key: 0be4ea9807
xp: 100
```

fdsfsdfsfsaf fdsfsdfsfsaf fdsfsdfsfsaf fdsfsdfsfsaf

`@instructions`
dsafdasfdfa fdsfsdfsfsaf fdsfsdfsfsaf fdsfsdfsfsaf

`@hint`
fdsasdfffd fdsfsdfsfsaf fdsfsdfsfsaf fdsfsdfsfsaf

`@pre_exercise_code`
```{r}
a <- c(-, 2, 3, 4, 5, 6)
```

`@sample_code`
```{r}

```

`@solution`
```{r}
a <- c(1, 2, 3, 4, 5, 6)
```

`@sct`
```{r}
ex() %>% check_expr("a[c(2, 4)]") %>% check_result() %>% check_equal()
```
