---
title: 'Motivation for regression'
description: 'Chapter description goes here.'
free_preview: true
---

## gfdgd

```yaml
type: NormalExercise
key: 0be4ea9807
xp: 100
```

fdsfsdfsfsaf fdsfsdfsfsaf fdsfsdfsfsaf fdsfsdfsfsaf

`@instructions`
dsafdasfdfa fdsfsdfsfsaf fdsfsdfsfsaf fdsfsdfsfsaf

`@hint`
fdsasdfffd fdsfsdfsfsaf fdsfsdfsfsaf fdsfsdfsfsaf

`@pre_exercise_code`
```{r}
plot(cars[---])
```

`@sample_code`
```{r}

```

`@solution`
```{r}
plot(cars[1,])
```

`@sct`
```{r}

```
