---
title: 'Motivation for regression'
description: 'Chapter description goes here.'
free_preview: true
attachments:
    slides_link: 'https://s3.amazonaws.com/assets.datacamp.com/production/course_17305/slides/chapter1.pdf'
---

## gfdgd

```yaml
type: NormalExercise
key: 0be4ea9807
xp: 100
```

fdsfsdfsfsaf fdsfsdfsfsaf fdsfsdfsfsaf fdsfsdfsfsaf

`@instructions`
dsafdasfdfa fdsfsdfsfsaf fdsfsdfsfsaf fdsfsdfsfsaf

`@hint`
fdsasdfffd fdsfsdfsfsaf fdsfsdfsfsaf fdsfsdfsfsaf

`@pre_exercise_code`
```{r}
my_fun <- function() { print('----') }
```

`@sample_code`
```{r}
my_fun <- function() { print('hello') }
```

`@solution`
```{r}
my_fun <- function() { print('hello') }
```

`@sct`
```{r}
ex() %>% check_error()
```

---

## motivation

```yaml
type: VideoExercise
key: 88fc46f18a
xp: 50
```

`@projector_key`
93919c62df44dc7ce29df74b631f26e5
