---
title: 'Motivation for regression'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.This is an example exercise. This is an example exercise.

`@instructions`
This is an example exercise.This is an example exercise.This is an example exercise.

`@hint`
This is an example exercise.

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 95e4633548
xp: 50
```

`@projector_key`
9efab82e37f8315af097770632b7758b
