---
title: 'Motivation for regression'
description: 'Chapter description goes here.'
free_preview: true
---

## gfdgd

```yaml
type: NormalExercise
key: 0be4ea9807
xp: 100
```

fdsfsdfsfsaf fdsfsdfsfsaf fdsfsdfsfsaf fdsfsdfsfsaf

`@instructions`
dsafdasfdfa fdsfsdfsfsaf fdsfsdfsfsaf fdsfsdfsfsaf

`@hint`
fdsasdfffd fdsfsdfsfsaf fdsfsdfsfsaf fdsfsdfsfsaf

`@pre_exercise_code`
```{r}
my_fun <- function() { print('----') }
```

`@sample_code`
```{r}
my_fun <- function() { print('hello') }
```

`@solution`
```{r}
my_fun <- function() { print('hello') }
```

`@sct`
```{r}
ex() %>% check_expr("my_fun()") %>% check_output() %>% check_equal()
```
