---
title: fdsfsfs
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise
key: 0beb5b0a4c
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}
my_fun <- function() { print('----') }
```

`@sample_code`
```{r}
dasdad
```

`@solution`
```{r}
my_fun <- function() { print('----') }
```

`@sct`
```{r}
ex() %>% check_error()
```
