---
title: fdsfsfs
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise
key: 0beb5b0a4c
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}
a <- c(-, 2, 3, 4, 5, 6)
```

`@sample_code`
```{r}

```

`@solution`
```{r}
a <- c(-, 2, 3, 4, 5, 6)
```

`@sct`
```{r}

```
